---
id: 68b7f31e79558b2bba0c285a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you'll learn about encapsulation, properties, and validation by building a salary tracking system for employees.

Start by creating an empty `Employee` class. This will be the foundation for implementing object-oriented programming concepts like private attributes and property decorators.

# --hints--

You should create an `Employee` class.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class('Employee')
    `)
})
```

The `Employee` class should be empty for now.

```js
({
    test: () => runPython(`
    class_body = _Node(_code).find_class('Employee').find_body()
    assert str(class_body).strip() == 'pass' or class_body == _Node()
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--


--fcc-editable-region--
```
