---
id: 68b7f5adda89e93c0cb8d1e0
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Now let's see what happens when we assign an invalid value. This demonstrates why we need data protection.

Try assigning the number `123` to `employee.name` instead of a string. This should work (unfortunately), showing that public attributes don't provide any data validation.

Print the employee again to see that Python allows this invalid assignment.

# --hints--

You should assign the number `123` to `employee.name`.

```js
({
    test: () => runPython(`
    assert 'employee.name' in _code and '123' in _code
    `)
})
```

The assignment should succeed, demonstrating the problem with public attributes.

```js
({
    test: () => runPython(`
    assert employee.name == 123
    `)
})
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level
    
    def __str__(self):
        return f'{self.name}: {self.level}'

employee = Employee('John Doe', 'junior')
print(employee)
employee.name = "Jane Doe"
print(employee)
--fcc-editable-region--

--fcc-editable-region--
```
