---
id: 68b7f5adda89e93c0cb8d1e2
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now that we've commented out the problematic code, let's make the `name` attribute private. In Python, you make an attribute private by prefixing it with double underscores (`__`).

Change `self.name = name` to `self.__name = name` in the `__init__` method. This creates a private attribute that cannot be accessed directly from outside the class.

# --hints--

You should change `self.name = name` to `self.__name = name` in the constructor.

```js
({
    test: () => runPython(`
    init_func = _Node(_code).find_class('Employee').find_function('__init__')
    assert init_func.has_variable('self.__name')
    `)
})
```

You should not have `self.name = name` in the constructor anymore.

```js
({
    test: () => runPython(`
    init_func = _Node(_code).find_class('Employee').find_function('__init__')
    name_vars = init_func.find_variables('self.name')
    assert len(name_vars) == 0
    `)
})
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
--fcc-editable-region--
        self.name = name
--fcc-editable-region--
        self.level = level
    
    def __str__(self):
        return f'{self.name}: {self.level}'

# employee = Employee('John Doe', 'junior')
# print(employee)
# employee.name = "Jane Doe"
# print(employee)
# employee.name = 123
# print(employee)
```
