---
id: 68b7f5adda89e93c0cb8d1e1
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Now we've seen the problem with public attributes - they can be changed to invalid values. Before we fix this by making attributes private, we need to temporarily comment out the code that uses the Employee class.

Comment out the lines that create and use the employee instance by adding `#` at the beginning of each line. This will prevent errors when we make changes to the class structure.

# --hints--

You should comment out the line that creates the employee instance.

```js
({
    test: () => runPython(`
    assert '#employee = Employee(' in _code or '# employee = Employee(' in _code
    `)
})
```

You should comment out all the print statements.

```js
({
    test: () => runPython(`
    lines = _code.split('\\n')
    print_lines = [line for line in lines if 'print(' in line and not line.strip().startswith('#')]
    assert len(print_lines) == 0
    `)
})
```

You should comment out the employee attribute assignments.

```js
({
    test: () => runPython(`
    assert '#employee.name' in _code or '# employee.name' in _code
    `)
})
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level
    
    def __str__(self):
        return f'{self.name}: {self.level}'

employee = Employee('John Doe', 'junior')
print(employee)
employee.name = "Jane Doe"
print(employee)
--fcc-editable-region--
employee.name = 123
print(employee)
--fcc-editable-region--
```
