---
id: 68b7f5adda89e93c0cb8d1e3
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Now make the `level` attribute private as well. Change `self.level = level` to `self.__level = level` in the `__init__` method.

This applies the same encapsulation principle to the level attribute.

# --hints--

You should change `self.level = level` to `self.__level = level` in the constructor.

```js
({
    test: () => runPython(`
    init_func = _Node(_code).find_class('Employee').find_function('__init__')
    assert init_func.has_variable('self.__level')
    `)
})
```

You should not have `self.level = level` in the constructor anymore.

```js
({
    test: () => runPython(`
    init_func = _Node(_code).find_class('Employee').find_function('__init__')
    level_vars = init_func.find_variables('self.level')
    assert len(level_vars) == 0
    `)
})
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.__name = name
--fcc-editable-region--
        self.level = level
--fcc-editable-region--
    
    def __str__(self):
        return f'{self.name}: {self.level}'

# employee = Employee('John Doe', 'junior')
# print(employee)
# employee.name = "Jane Doe"
# print(employee)
# employee.name = 123
# print(employee)
```
