---
id: 68b7f31e79558b2bba0c285a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you'll practice encapsulation, properties, and validation by building a salary tracking system for employees.

Start by creating an empty `Employee` class.

# --hints--

You should create an `Employee` class.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class('Employee')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
